<subject>
Assignment How-To
</subject>

<details>
    
**Do Not Talk During Quizzes or Exams**

Do not talk at all once an exam has begun. You may talk again once you leave the room after the quiz or exam. The one exception to this rule is if you need to ask the exam proctor a question. Any talking during a quiz or exam qualifies the student for disciplinary action.

**Naming Conventions**

When naming your files for upload, you must follow the format below:

    <uni>_<assignment>_<details [optional]>.<extension>

For example, if I needed to hand in HW 0, any of the below formats would be sufficient for a file name:

- pl2648_hw0.ipynb
- pl2648_hw0.txt
- pl2648_hw0.sh
- pl2648_hw0_all_in_one.txt
- Pl2648_hw0_bash_program.sh
- Pl2648_quiz1.ipynb

This naming format allows for autograding of all assignments. If your files are not named with this format, you should expect a grade of zero for the assignment.

Courseworks may rename your file to something like `pl2648_hw0-1.ipynb` if you resubmit your assignment. This is perfectly fine.

**What Format To Submit In**

Most homework and quizzes are in Jupyter notebooks. Unless specified otherwise, please download your work as an `.ipynb` file from your local machine and upload it to courseworks.


**Grading**

Possible points on late assignments are deducted by 50% for each day they are late. For example, if you get 80% of the total possible credits on a homework but hand in that homework a day late, you would get 40%. Assignments two days late get zero points.

Once solutions are posted and graded assignments are handed back, students have 1 week to bring their grading discrepancies to a CA for consideration of possible grading errors. 

Because grading is automated, please delete (or comment out) the `raise NotImplmeneted` code before attempting a problem. 

Empty un-editable cells in an assignment are there for a purpose. They will be filled with tests by the automatic grader. Please do not attempt to remove them.

**Getting Help**

Asking for help is a great way to increase your chance of success. However there are some rules. When asking for help (especially from a fellow student), *you can show your helper your code but you can not view theirs*. You work needs to be your own. You can not post screenshots of your current work to Piazza or other tools used for getting help.

If you need to reach out to a CA for help, please do so via Piazza and not via email. Answers given via Piazza will help you as well as other students. Thus, emails will always have a lower priority for response than Piazza questions. If you do email the CA, please make a note of what section you are in. This helps us identify you in courseworks faster. 

Finally, if you do not get a repsonse from a CA within 48 hours, you may email the professor.

**Multiple Choice**

If the question is multiple choice, you will be given several options to choose from and your function will need to return **one** of those options **verbatim** as a string.

For example: 

Which of the following animals bark?

- dogs
- cats
- fish
- trees

A correct answer would be structured in the following way:

```python
def question_animals():
    return 'dogs'
```

You answer will be stripped of left and right white space and lowercased before comparison to the correct answer during grading.
</details>

For the following questions, assume you have the following tables of information in the database `nmh`. 

- table name: `exhibits`

id (INT) | name (VARCHAR) | start_date (DATE) | end_date (DATE) | curator_id (INT) [FK to curator.id]
-|-|-|-|-
3| Free The Fishes | 2018-01-01 | 2018-06-30 | 5
17| Space, What Lies Above | 2018-02-01 | 2018-05-30 | 11
23| Bears Bears Bears | 2018-02-14 | 2018-02-24 | 5
46| Humans? Aliens? | 2019-03-14 | 2019-10-21 | 11


- table name: `curators`

id (INT) | name (VARCHAR) | bio (TEXT)
-|-|-
5| Rebecca Votea | Esteemed naturalist
11| Simon Strauss | Space man
71| Rick Sanchez | Grandfather


SQL solutions will be graded on their simplicity as well as on whether or not they would return the correct answer from the database.

**Q**: Write a function that returns a string containing the SQL necessary to create two tables above.

These tables are necessary for future problems.

[0 points]

<!-- grade_id: cell-518e4070a5c1b4d1, points: 0, grade: false, solution: true -->
```python
def create_tables_sql() -> str:
    return """
    CREATE TABLE exhibits (
        id integer PRIMARY KEY,
        name varchar(255),
        start_date date,
        end_date date,
        curator_id int,
        FOREIGN KEY(curator_id) REFERENCES curators(id)
    );

    CREATE TABLE curators (
        id integer PRIMARY KEY,
        name varchar(255),
        bio text
    );

    """
        
```

<!-- grade_id: cell-f0b2a4f52364ca3e, points: 0, grade: true, solution: false -->
```python
### BEGIN TESTS
import sqlite3
conn = sqlite3.connect(':memory:')

c = conn.cursor()

c.executescript(create_tables_sql())

conn.commit()
### END TESTS
```

**Q**: Write a function that returns a string containing the SQL necessary to add the data from the two tables above.

This data is necessary for future problems.

[0 points]

<!-- grade_id: cell-dd95d3a6b4b0d7dd, points: 0, grade: false, solution: true -->
```python
def insert_data_sql() -> str:
    return """
    INSERT INTO exhibits VALUES (3, 'Free The Fishes', '2018-01-01', '2018-06-30', 5);
    INSERT INTO exhibits VALUES (17, 'Space, What Lies Above', '2018-02-01', '2018-05-30', 11);
    INSERT INTO exhibits VALUES (23, 'Bears Bears Bears', '2018-02-14', '2018-02-24', 5);
    INSERT INTO exhibits VALUES (46, 'Humans? Aliens?', '2019-03-14', '2019-10-21', 11);
    INSERT INTO curators VALUES (5, 'Rebecca Votea', 'Esteemed naturalist');
    INSERT INTO curators VALUES (11, 'Simon Strauss', 'Space man');
    INSERT INTO curators VALUES (71,'Rick Sanchez', 'Grandfather');"""
```

<!-- grade_id: cell-3de325c72a925e5e, points: 0, grade: true, solution: false -->
```python
### BEGIN TESTS
c.executescript(insert_data_sql())
conn.commit()
### END TESTS
```

**Q**: Create a function that will return the SQL to find all the names of all the curators ordered by their id. IDs should be in descending order. Do not use an alias for the table name.

[1 point]

<!-- grade_id: cell-47e83ca712cccece, points: 0, grade: false, solution: true -->
```python
def get_sql():
    return """
    SELECT name FROM curators ORDER BY id DESC;
    """
```

<!-- grade_id: cell-8c584e66fca1031f, points: 1, grade: true, solution: false -->
```python
### BEGIN TESTS
c.execute(get_sql())
result = c.fetchall()
assert result == [('Rick Sanchez',), ('Simon Strauss',), ('Rebecca Votea',)]
### BEGIN TESTS
```

**Q**: Create a function that will return the SQL to find all the names of all exhibts and the names of their associated curators. If an exhibit does not have a curator, the exhibit should not be in the result. Even if a curator does not have an exhibit, the curator should be present in the final result. The first letter of the table name should be used as its alias.

For example,

```
Rebecca Votea, Free The Fishes
Rebecca Votea, Bears Bears Bears
Rick Sanchez, 
...
```

[1 point]

<!-- grade_id: cell-dd55b2d154d1f0b8, points: 0, grade: false, solution: true -->
```python
def get_sql():
    return """
    SELECT c.name, e.name
    FROM curators c
    LEFT JOIN exhibits e ON c.id=e.curator_id
    """
```

```python
c.execute(get_sql())
set(c.fetchall())
```

Output:
```
{('Rebecca Votea', 'Bears Bears Bears'),
 ('Rebecca Votea', 'Free The Fishes'),
 ('Rick Sanchez', None),
 ('Simon Strauss', 'Humans? Aliens?'),
 ('Simon Strauss', 'Space, What Lies Above')}
```

<!-- grade_id: cell-2ebddb0b10914442, points: 1, grade: true, solution: false -->
```python
### BEGIN TESTS
c.execute(get_sql())
result = set(c.fetchall())
expected = {
    ('Rebecca Votea', 'Bears Bears Bears'),
    ('Rebecca Votea', 'Free The Fishes'),
    ('Rick Sanchez', None),
    ('Simon Strauss', 'Humans? Aliens?'),
    ('Simon Strauss', 'Space, What Lies Above')
}
assert result == expected
### BEGIN TESTS
```

**Q**: Create a function that will return the SQL to find the number of exhibits for all curators with a name that starts with "R" (case-insensitive). The output of running the SQL should be two fields (name, count). 

Even if a curator does not have an exhibit, the curator should be present in the final result. The first letter of the table name should be used as its alias. You can not assume curators names are distinct.

[1 point]

<!-- grade_id: cell-0381ce7f2909db88, points: 0, grade: false, solution: true -->
```python
def get_sql():
    return """
    SELECT c.name, COUNT(e.name)
    FROM curators c
    LEFT JOIN exhibits e ON c.id=e.curator_id
    WHERE c.name LIKE 'R%'
    GROUP BY c.name
    """
c.execute(get_sql())
result = set(c.fetchall())
result
```

Output:
```
{('Rebecca Votea', 2), ('Rick Sanchez', 0)}
```

<!-- grade_id: cell-156cced94b512b38, points: 1, grade: true, solution: false -->
```python
### BEGIN TESTS
c.execute(get_sql())
result = set(c.fetchall())
expected = {
    ('Rebecca Votea', 2), 
    ('Rick Sanchez', 0),
}
assert result == expected
### BEGIN TESTS
```

**Q**: Create a function that will return the SQL to find the exhibits with overlapping runs (overlapping intervals of time bound by start date and end date (inclusive)). The result should include only one instance of each overlapping pair. The exhibit with the lower valued ID should be listed as `exhibit1` in the result and the other with `exhibit2`.

The result should be:

```
        exhibit1        |        exhibit2        
------------------------+------------------------
 Free The Fishes        | Space, What Lies Above
 Free The Fishes        | Bears Bears Bears
 Space, What Lies Above | Bears Bears Bears
```

The first letter of the table name should be used as its alias. If multiple instances of the same table are necessary, append a number to the end of the alias. The number sequence appended to the alias should start at 1.

[1 point]

<!-- grade_id: cell-31cbf33af0444f1a, points: 0, grade: false, solution: true -->
```python
def get_sql():
    # YOUR CODE HERE
    return """
    SELECT e1.name, e2.name
    FROM exhibits e1, exhibits e2
    WHERE
    /*Check for overlap: */
    (e1.end_date >= e2.start_date) AND (e1.start_date <= e2.end_date)
    
    /*Remove self maps:*/
    AND (e1.name <> e2.name) 
    
    /*Sort columns by id*/
    AND (e1.id < e2.id)
    """
    
c.execute(get_sql())
result = set(c.fetchall())
result = tuple(map(set, result))
result
```

Output:
```
({'Bears Bears Bears', 'Free The Fishes'},
 {'Free The Fishes', 'Space, What Lies Above'},
 {'Bears Bears Bears', 'Space, What Lies Above'})
```

<!-- grade_id: cell-068a6237adc2e06c, points: 1, grade: true, solution: false -->
```python
### BEGIN TESTS
c.execute(get_sql())
result = set(c.fetchall())
result = tuple(map(set, result))

expected = (
    {'Bears Bears Bears', 'Space, What Lies Above'},
    {'Bears Bears Bears', 'Free The Fishes'}, 
    {'Free The Fishes', 'Space, What Lies Above'}
)

for set_ in expected:
    assert set_ in result

assert len(result) == 3
### BEGIN TESTS
```

<!-- grade_id: cell-1249017f35452bbc, points: 0, grade: true, solution: false -->
```python

```

```python

```