**Q**: Build a function to return a regex expression. The expression, when passed to the appropriate `re` function, should find all of the dates of the format YYYY-MM-DD in the in following string.

```python
>>> string = """
2010-04-29 1984 ABC-123!
Feb 18th, 2018. 
888-27-8949 2048-07-07
"""

>>> pattern = get_pattern()
>>> func = get_function()
>>> func(pattern, string)
['2010-04-29', '2048-07-07']
```

[1 point]

<!-- grade_id: cell-191bb99330f39907, points: 0, grade: false, solution: true -->
```python
import re

def get_pattern():
    return re.compile(r'(?<!\d)\d{4}-\d{2}-\d{2}(?!\d)')

def get_function():
    return re.findall
```

<!-- grade_id: cell-7c5fc12bf1bedd77, points: 0.25, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert get_function() == re.findall
### END HIDDEN TESTS
```

<!-- grade_id: cell-37f7551c7035fb27, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
string = """
2010-04-29 1984 ABC-123!
Feb 18th, 2018. 
888-27-8949 2048-07-07
12048-07-074
"""

pattern = get_pattern()
func = get_function()
result = func(pattern, string)

assert result == ['2010-04-29', '2048-07-07'], result
### END TESTS
```

<!-- grade_id: cell-7719cc19d23054ae, points: 0.25, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
string = """
2-1-1
00-00-00
000-000-000
0000-00-00
12452-02-12
12452-02-1
"""

pattern = get_pattern()
func = get_function()
result = func(pattern, string)

assert result == ['0000-00-00'], result
### END HIDDEN TESTS
```

**Q**: Write a function that takes a string like the following and returns a match object if a match is possible. 

    '2018 Mar 24 02:41:00 AM'

The match object should have the following named groups:

    - year
    - month_name
    - day
    - hour
    - minute
    - second
    - meridiem
    
For example, the output of ...

```python
match = pattern.match('2018 Mar 24 02:41:00 AM')
print('MATCH', match, sep=': ')
print('GROUPS', match.groups(), sep=': ')

for name in ('year', 'month_name', 'day', 'hour', 'minute', 'second', 'meridiem'):
    print(name.upper(), match.group(name), sep=': ')
 ```
where `pattern` is the compiled regex expression should be:

    MATCH: <_sre.SRE_Match object; span=(0, 23), match='2018 Mar 24 02:41:00 AM'>
    GROUPS: ('2018', 'Mar', '24', '02', '41', '00', 'AM')
    YEAR: 2018
    MONTH_NAME: Mar
    DAY: 24
    HOUR: 02
    MINUTE: 41
    SECOND: 00
    MERIDIEM: AM
    

[2 points]

<!-- grade_id: cell-be554692bfc729a9, points: 0, grade: false, solution: true -->
```python
import re

def get_match_object(string):
    pattern = re.compile(r'(?P<year>\d{4})\s(?P<month_name>[a-zA-Z]{3})\s(?P<day>\d{2})\s'\
                        +r'(?P<hour>\d{2}):(?P<minute>\d{2}):(?P<second>\d{2})\s(?P<meridiem>[A-Z]{2})')
    return re.match(pattern, string)

    # YOUR CODE HERE
    #raise NotImplementedError()
```

<!-- grade_id: cell-20885ead6e5a1fac, points: 0.25, grade: true, solution: false -->
```python
### BEGIN TESTS
match = get_match_object('2018 Mar 24 02:41:00 AM')
assert match.group('year') == '2018'
### END TESTS
```

<!-- grade_id: cell-d65adb32954b254a, points: 0.25, grade: true, solution: false -->
```python
### BEGIN TESTS
match = get_match_object('2018 Mar 24 02:41:00 AM')
assert match.group('month_name') == 'Mar'
### END TESTS
```

<!-- grade_id: cell-64860e1dd88fbc1d, points: 0.25, grade: true, solution: false -->
```python
### BEGIN TESTS
match = get_match_object('2018 Mar 24 02:41:00 AM')
assert match.group('day') == '24'
### END TESTS
```

<!-- grade_id: cell-47661529f963e201, points: 0.25, grade: true, solution: false -->
```python
### BEGIN TESTS
match = get_match_object('2018 Mar 24 02:41:00 AM')
assert match.group('hour') == '02'
### END TESTS
```

<!-- grade_id: cell-2e203308d6349075, points: 0.25, grade: true, solution: false -->
```python
### BEGIN TESTS
match = get_match_object('2018 Mar 24 02:41:00 AM')
assert match.group('minute') == '41'
### END TESTS
```

<!-- grade_id: cell-51337e5d0dacb4ad, points: 0.25, grade: true, solution: false -->
```python
### BEGIN TESTS
match = get_match_object('2018 Mar 24 02:41:00 AM')
assert match.group('second') == '00'
### END TESTS
```

<!-- grade_id: cell-d49f5d2076486460, points: 0.25, grade: true, solution: false -->
```python
### BEGIN TESTS
match = get_match_object('2018 Mar 24 02:41:00 AM')
assert match.group('meridiem') == 'AM'
### END TESTS
```

<!-- grade_id: cell-d5f2c2c1e38f1be6, points: 0.25, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
match = get_match_object('1999 Dec 31 11:59:59 PM')
groups = ('year', 'month_name', 'day', 'hour', 'minute', 'second', 'meridiem')
assert [match.group(name) for name in groups] == ['1999', 'Dec', '31', '11', '59', '59', 'PM']
### END HIDDEN TESTS
```

**Q**: NYC publishes 311 call reports for public review. An example slice of that data has been made available at the following link:

https://data.cityofnewyork.us/Social-Services/311-11-19-2018-/6ir3-7kwf/data

BUT ... please download a conditioned CSV of this data from courseworks. It will be named "311__11_19_2018.csv".

Using the `re` module, find the number at the beginning of each "Incident Address". 

For "17 WEST 74 STREET" the number would be 17, for "88-28 74 PLACE" the number would be 88-28. Ignore addresses that are empty strings. If the street number is 123A, consider the number to be 123. An address must start with a number to be considered; so addresses like "-123" should not be considered. 

If the number is a hyphenated number, peform the subtraction (ie. 88-28 = 60). Sum up all of the results of the hyphenated addresses. (ie. 60 + (24-11) + ...). 

Sum up all the un-hyphenated addresses.

Return the pair as a tuple from the function provided. The function should take the absolute path to the CSV file.

[2 points]

```python

uh = 0
h = 0 
with open('311__11_19_2018_.csv', 'r') as f:
        pattern = re.compile(r'^\d[\d-]*')
        #hyph = []
        #unhyph = []

        
        for line in f.readlines():
            address = re.split(r',', line)[9]
            if address:
                res = re.search(pattern, address) 
                if res:
                    number = res.group(0)
                    if '-' in number:
                        sp = [int(v) for v in number.split('-')]
                        uh += (sp[0]-sp[1])
                    else:
                        h += int(number)
                        
uh, h
```

Output:
```
(5527, 288085)
```

<!-- grade_id: cell-0b8c920f2ed2ebfc, points: 0, grade: false, solution: true -->
```python
def get_address_sums(path: str) -> (int, int):
    with open(path) as f:
        pattern = re.compile(r'^\d[\d-]*')
        hyph = []
        unhyph = []


        for line in f.readlines():
            address = re.split(r',', line)[9]

            #For case tests:
            #address = '123A'
            if address:
                res = re.search(pattern, address) 
                if res:
                    number = res.group(0)
                    if '-' in number:
                        sp = [int(v) for v in number.split('-')]
                        
                        #subtract all
                        first = sp[0]
                        rest = sp[1:]
                        for v in rest:
                            first -= v
                        hyph.append(first)
                    else:
                        unhyph.append(int(number))


                    #print(address, ' | ', number)
    result = (sum(hyph), sum(unhyph))
    return result
```

<!-- grade_id: cell-695afe2438cf8f12, points: 0, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
unhyphenated, hyphenated = get_address_sums('311__11_19_2018_.csv')
### END HIDDEN TESTS
```

<!-- grade_id: cell-7ab2046a0323a2a7, points: 1, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert unhyphenated == 278084
### END HIDDEN TESTS
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-18-5dfcfa7ec178> in <module>
      1 ### BEGIN HIDDEN TESTS
----> 2 assert unhyphenated == 278084
      3 ### END HIDDEN TESTS

AssertionError:
```

<!-- grade_id: cell-3d7a0349af2c6f3f, points: 1, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert hyphenated == 5527
### END HIDDEN TESTS
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-19-a38d64d0d800> in <module>
      1 ### BEGIN HIDDEN TESTS
----> 2 assert hyphenated == 5527
      3 ### END HIDDEN TESTS

AssertionError:
```

```python
hyphenated
```

Output:
```
288085
```

```python

```