<subject>
Assignment How-To
</subject>

<details>

**Naming Conventions**

When naming your files for upload, you must follow the format below:

    <uni>_<assignment>_<details [optional]>.<extension>

For example, if I needed to hand in HW 0, any of the below formats would be sufficient for a file name:

- pl2648_hw0.ipynb
- pl2648_hw0.txt
- pl2648_hw0.sh
- pl2648_hw0_all_in_one.txt
- Pl2648_hw0_bash_program.sh

This naming format allows for autograding of all assignments. If your files are not named with this format, you should expect a grade of zero for the assignment.

**Grading**

Possible points on late homeworks are deducted by 50% for each day they are late. For example, if you get 80% of the total possible credits on a homework but hand in that homework a day late, you would get 40%. Assignments two days late get zero points.

Once solutions are posted and graded assignments are handed back, students have 1 week to bring their grading discrepancies to a CA for consideration of possible grading errors. 

**Getting Help**

Asking for help is a great way to increase your chance of success. However there are some rules. When asking for help (especially from a fellow student), *you can show your helper your code but you can not view theirs*. You work needs to be your own. 

If you need to reach out to a CA or the Professor for help, please do so via Piazza and not via email. Answers given via Piazza will help you as well as other students. Thus, emails will always have a lower priority for response than Piazza questions. If you do email the professor or TA, please make a note of what section you are in. This helps us identify you in courseworks faster.

**Multiple Choice**

If the question is multiple choice, you will be given several options to choose from and your function will need to return **one** of those options **verbatim** as a string.

For example: 

Which of the following animals bark?

- dogs
- cats
- fish
- trees

A correct answer would be structured in the following way:

```python
def question_animals():
    return 'dogs'
```

You answer will be stripped of left and right white space and lowercased before comparison to the correct answer during grading.
</details>

**Q**: Write a function `verify` that takes a string as user input (via arguments! not via the built-in function `input`) and returns the string if it contains a number, an uppercase letter, and a lowercase letter. The function should raise a `ValueError` otherwise. The `ValueError` should have the exact message `'Must contain a number, an uppercase letter, and a lowercase letter'`.

[1 point]

<!-- grade_id: cell-51ed750f23e67ce9, points: 0, grade: false, solution: true -->
```python
def verify(string: str) -> str:
    bool_list = []
    for char in string:
        bool_list.append(char.isupper())
    has_upper = (True in bool_list)
    bool_list = []
    for char in string:
        bool_list.append(char.islower())
    has_lower = (True in bool_list)
    bool_list = []
    for char in string:
        bool_list.append(char.isdigit())
    has_digit = (True in bool_list)
    try:
        assert(has_upper and has_lower and has_digit)
        return string
    
    except:
        raise ValueError('Must contain a number, an uppercase letter, and a lowercase letter')
        
    # YOUR CODE HERE
    #raise NotImplementedError()
```

<!-- grade_id: cell-8558a10afc63169f, points: 0, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
msg = 'Must contain a number, an uppercase letter, and a lowercase letter'
### END HIDDEN TESTS
```

<!-- grade_id: cell-43a9d60751a75846, points: 0.25, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
try:
    verify('ab1')
    assert False, 'Should not pass'
except ValueError as e:
    assert str(e) == msg
### END HIDDEN TESTS    
```

<!-- grade_id: cell-286dd59ff59f783e, points: 0.25, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
try:
    verify('AB1')
    assert False, 'Should not pass'
except ValueError as e:
    assert str(e) == msg
### END HIDDEN TESTS   
```

<!-- grade_id: cell-b5e1ec457b072396, points: 0.25, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
try:
    verify('Abc')
    assert False, 'Should not pass'
except ValueError as e:
    assert str(e) == msg
### END HIDDEN TESTS   
```

<!-- grade_id: cell-2b369f2e4501cc5b, points: 0.25, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
try:
    verify('Ab1')
    assert True
except ValueError as e:
    assert False, 'Should pass'
### END HIDDEN TESTS   
```

**Q**: Crearte a predicate function (aka. a function that returns `True` or `False` to qualify or categorize its arguments) named `is_valid_divisor`. 

Your function should return `True` if a value is a valid option as the divisor in a mathematical division operation. 

For example, `2` would be a valid divisor while `"hi"` would not.

Your function will be passed to `filter` as its first argument to test an interable of values. 

[1 point]

```python
type(int)
```

Output:
```
type
```

<!-- grade_id: cell-4feca3e1d1ef9168, points: 0, grade: false, solution: true -->
```python
def is_valid_divisor(value):
    try:
        test = 1/value
        return True
    except:
        return False
    
    # YOUR CODE HERE
    #raise NotImplementedError()
```

<!-- grade_id: cell-e2ba8fdb61141305, points: 0.25, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert is_valid_divisor(1)
### END HIDDEN TESTS
```

<!-- grade_id: cell-266f74082d64c1b5, points: 0.25, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert is_valid_divisor('1') == False
### END HIDDEN TESTS
```

<!-- grade_id: cell-a79d049474e317b1, points: 0.25, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert is_valid_divisor(0) == False
### END HIDDEN TESTS
```

<!-- grade_id: cell-661d9daae76c9226, points: 0.25, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert is_valid_divisor(None) == False
### END HIDDEN TESTS
```

```python

```