<subject>
Homework How-To
</subject>

<details>

**Naming Conventions**

When naming your files for upload, you must follow the format below:

    <uni>_<assignment>_<details [optional]>.<extension>

For example, if I needed to hand in HW 0, any of the below formats would be sufficient for a file name:

- pl2648_hw0.txt
- pl2648_hw0.sh
- pl2648_hw0_all_in_one.txt
- Pl2648_hw0_bash_program.sh

This naming format allows for autograding of all assignments. If your files are not named with this format, you should expect a grade of zero for the assignment.

**Grading**

Possible points on late homeworks are deducted by 50% for each day they are late. For example, if you get 80% of the total possible credits on a homework but hand in that homework a day late, you would get 40%. Assignments two days late get zero points.

Once solutions are posted and graded assignments are handed back, students have 1 week to bring their grading discrepancies to a CA for consideration of possible grading errors. 

**Getting Help**

Asking for help is a great way to increase your chance of success. However there are some rules. When asking for help (especially from a fellow student), *you can show your helper your code but you can not view theirs*. You work needs to be your own. 

If you need to reach out to a CA or the Professor for help, please do so via Piazza and not via email. Answers given via Piazza will help you as well as other students. Thus, emails will always have a lower priority for response than Piazza questions. If you do email the professor or TA, please make a note of what section you are in. This helps us identify you in courseworks faster.

**Multiple Choice**

If the question is multiple choice, you will be given several options to choose from and your function will need to return **one** of those options **verbatim** as a string.

For example: 

Which of the following animals bark?

- dogs
- cats
- fish
- trees

A correct answer would be structured in the following way:

```python
def question_animals():
    return 'dogs'
```

You answer will be stripped of left and right white space and lowercased before comparison to the correct answer during grading.
</details>

**Q**: Write a script that appends 250,000 items to a an empty list; the item can be anything so the number 1 would do. Time how long it takes (with a stop watch on the Internet or your phone). Write a script that prepends 250,000 items to an empty list and time how long it takes. *DO NOT PRINT THE RESULT OR EACH VALUE ADDED TO THESE LISTS. Please clear your cell's output of the 250,000 items before handing in your assignment; scrolling through 250,000 items is a lot of work!*

What script takes longer and why? No more than 50 words. 

If either script runs for more than a minute, stop it and assume that it ran infinitely long.

[1 point]

<!-- grade_id: cell-670e7107d58db7bf, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
list_1=[]
for element in range(0,250000):
    list_1.append(1)
#raise NotImplementedError()
```

<!-- grade_id: cell-c1f537f3db59f76e, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
list_1=[]
for element in range(0,250000):
    list_1.insert(0,1)
#raise NotImplementedError()
```

<!-- grade_id: cell-91f1edcf7c372b7c, points: 1, grade: true, solution: true -->
The prepend script takes longer time. This is because for insert(0,1) the pointer returns to th first position every single time when we call the function.

**Q**: Create a function `evens` that returns all the even numbers between 0 and a given positive number x. Zero should be included in the result but x should not be included unless x is zero. Any data structure that is iterable will be sufficient for returning these values. 

[1 point]

<!-- grade_id: cell-f178c23df522f905, points: 0, grade: false, solution: true -->
```python
list_1=[]
def evens(x):
    if x == 0:
        list_1 = [0]
        return list_1
    elif x > 0: 
        for i in range(x):
            if i%2==0:
                list_1.append(i)
        return list_1
print(evens(0))
    # YOUR CODE HERE
    #raise NotImplementedError()
```

Output:
```
[0]
```

<!-- grade_id: cell-80b80f97d1d635e1, points: 0.25, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-6d57f030024737d2, points: 0.25, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-ed8531a55caf7739, points: 0.25, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-f78b76614fd7bb17, points: 0.25, grade: true, solution: false -->
```python

```

**Q**: Implement Fizz-Buzz. Write a function `fizz_buzz` that returns the number given. 

But for multiples of two return "Fizz" instead of the number, for the multiples of three return "Buzz", and for the multiples of five return "Bang". For numbers which are multiples of both two and three return "FizzBuzz". For numbers which are multiples of two, three, and five return "FizzBuzzBang". Assume only positive numbers will be passed to the function. 

Order of precedence should be as follows (first has the highest precedence):

    - FizzBuzzBang
    - FizzBuzz
    - Bang
    - Buzz
    - Fizz


[2 points]

<!-- grade_id: cell-fda6d5a34cfe1fec, points: 0, grade: false, solution: true -->
```python
def fizz_buzz(element):
    if element%30==0:
        return "FizzBuzzBang"
    elif element%6==0:
        return "FizzBuzz"
    elif element%5==0:
        return "Bang"
    elif element%3==0:
        return "Buzz"
    elif element%2==0:
        return "Fizz"
    else: 
        return element
```

<!-- grade_id: cell-edc1cee03a574266, points: 0.25, grade: true, solution: false -->
```python
### BEGIN TESTS
assert fizz_buzz(1) == 1
### END TESTS
```

<!-- grade_id: cell-4f8d689ae39f86b7, points: 0.25, grade: true, solution: false -->
```python
### BEGIN TESTS
assert fizz_buzz(2) == 'Fizz'
### END TESTS
```

<!-- grade_id: cell-3243239057da9c5d, points: 0.25, grade: true, solution: false -->
```python
### BEGIN TESTS
assert fizz_buzz(3) == 'Buzz'
### END TESTS
```

<!-- grade_id: cell-f6dc4ad9f061fe39, points: 0.25, grade: true, solution: false -->
```python
### BEGIN TESTS
assert fizz_buzz(5) == 'Bang'
### END TESTS
```

<!-- grade_id: cell-946d26b610de6b73, points: 0.25, grade: true, solution: false -->
```python
### BEGIN TESTS
assert fizz_buzz(6) == 'FizzBuzz'
### END TESTS
```

<!-- grade_id: cell-f9b98e67123e3631, points: 0.25, grade: true, solution: false -->
```python
### BEGIN TESTS
assert fizz_buzz(30) == 'FizzBuzzBang'
### END TESTS
```

<!-- grade_id: cell-eb510034e9d6fb63, points: 0.5, grade: true, solution: false -->
```python

```

**Q**: Write a function `compress` that takes a string of characters, counts each occurrence of each character, and produces a compressed string with each character and its count. For example:

`'aaaabbbbcc'` would compress to `'a4b4c2'`

You can assume that the source string will be sorted such that like characters are adjacent. The result string should be sorted. You do not have to consider empty strings or infinite strings.

[2 points]

<!-- grade_id: cell-63c10b65c17caab0, points: 0, grade: false, solution: true -->
```python
set_1=set()
def compress(letter_string: str)-> str:
    
    all_freq = {}
    for i in letter_string: 
        if i in all_freq: 
            all_freq[i] += 1
        else: 
            all_freq[i] = 1
            
    x = ''.join('{}{}'.format(key, val) for key, val in sorted(all_freq.items()))
    return x    
compress('aaaabbbbcc')

    # YOUR CODE HERE
    #raise NotImplementedError()

```

Output:
```
'a4b4c2'
```

<!-- grade_id: cell-8d41bb3300810b31, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
assert compress('aaaabbbbcccc') == 'a4b4c4'
### END TESTS
```

<!-- grade_id: cell-7b5d55b9412941b1, points: 0.5, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-33ec73d3be0d6640, points: 1, grade: true, solution: false -->
```python

```

**Q**: Write a function `endpoints` that takes a list of numbers (eg. `[5, 10, 15, 20, 25]`) and returns a new list of only the first and last elements of the given list (eg. `[5, 25]`). The function should return an empty list if an empty list is passed in. The function should *not* use any variables besides the passed in argument `list_`.  

[1 point]

<!-- grade_id: cell-e06b186eadc3d5c3, points: 0, grade: false, solution: true -->
```python
a = [5, 10, 15, 20]
def endpoints(a) -> list:
    if a is []:
        return []
    else:
        return [a[0],a[-1]]
endpoints(a)
    
    # YOUR CODE HERE
    #raise NotImplementedError()
```

Output:
```
[5, 20]
```

<!-- grade_id: cell-34da01cd23f4b977, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
assert endpoints([5, 10, 15, 20, 25]) == [5, 25]
### END TESTS
```

<!-- grade_id: cell-f7e7bb26b30ed92a, points: 0.5, grade: true, solution: false -->
```python

```

**Q**: Write a function `common` that consumes two lists and that produces a list that contains only the common elements (without duplicates). Make sure your program works on two lists of different sizes and returns the common elements in sorted ascending order.


```python
a = [1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89]
b = [1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13]

# result
[1, 2, 3, 5, 8, 13]
```

[1 point]

<!-- grade_id: cell-1ab667ab32f392f9, points: 0, grade: false, solution: true -->
```python
a = {1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89}
b = {1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13}
def common(a, b):
    return (a & b)

common(a,b)
    
    # YOUR CODE HERE
    #raise NotImplementedError()
```

Output:
```
{1, 2, 3, 5, 8, 13}
```

<!-- grade_id: cell-8cd8aa2c0c722e82, points: 1, grade: true, solution: false -->
```python

```

**Q**: Create a function `get_lambda` that returns a `lambda` function that returns a tuple containing N instances of a given value.

```python
>>> f = get_lambda()
>>> f(4, True)
(True, True, True, True)
```

[1 point]

<!-- grade_id: cell-440167c0cf8ec924, points: 0, grade: false, solution: true -->
```python
def get_lambda():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-bf1ea39119b33fad, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
assert get_lambda()(4, True) == (True, True, True, True)
### END TESTS
```

<!-- grade_id: cell-d1f15d8c975bef1e, points: 0.5, grade: true, solution: false -->
```python

```

**Q**: Create a function `get_lambda` that returns a `lambda` function that takes a tuple and raises each element in the tuple to the second power.

```python
>>> f = get_lambda()
>>> f((1, 2, 3, 4))
(1, 4, 9, 16)
```
[1 point]

<!-- grade_id: cell-70896f7d7cc69cde, points: 0, grade: false, solution: true -->
```python
def get_lambda():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-2107d7a4f3fef8f6, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
assert get_lambda()((1, 2, 3, 4)) == (1, 4, 9, 16)
### END TESTS
```

<!-- grade_id: cell-7b39c52cfbb8fe3a, points: 0.5, grade: true, solution: false -->
```python

```

##### Extra Practice

Extra Practice problems are exactly that: extra practice. They are not worth any points or extra credit.

**Extra Practice**: Create a function that returns a list that contains itself as its one and only element. 

[0 points]

<!-- grade_id: cell-3af2459c61b6c472, points: 0, grade: false, solution: true -->
```python
def itself():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-6dbaa09e08641787, points: 0, grade: true, solution: false -->
```python
### BEGIN TESTS
list_ = itself()
assert list_ is list_[0]
### END TESTS
```

**Extra Practice**: Create a function that returns the first positive integer for which Python does not create a singleton. What about the negative integer closest to zero? 

[0 points]

<!-- grade_id: cell-6f599029a544dfcd, points: 0, grade: false, solution: true -->
```python
def first_non_singleton():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-8a115bd7fb268cc5, points: 0, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-f5077f7c620f08cb, points: 0, grade: false, solution: true -->
```python
def last_negative_non_singleton():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-4ea5e3cac5553b11, points: 0, grade: true, solution: false -->
```python

```