<subject>
Assignment How-To
</subject>

<details>

**Naming Conventions**

When naming your files for upload, you must follow the format below:

    <uni>_<assignment>_<details [optional]>.<extension>

For example, if I needed to hand in HW 0, any of the below formats would be sufficient for a file name:

- pl2648_hw0.ipynb
- pl2648_hw0.txt
- pl2648_hw0.sh
- pl2648_hw0_all_in_one.txt
- Pl2648_hw0_bash_program.sh

This naming format allows for autograding of all assignments. If your files are not named with this format, you should expect a grade of zero for the assignment.

**Grading**

Possible points on late homeworks are deducted by 50% for each day they are late. For example, if you get 80% of the total possible credits on a homework but hand in that homework a day late, you would get 40%. Assignments two days late get zero points.

Once solutions are posted and graded assignments are handed back, students have 1 week to bring their grading discrepancies to a CA for consideration of possible grading errors. 

**Getting Help**

Asking for help is a great way to increase your chance of success. However there are some rules. When asking for help (especially from a fellow student), *you can show your helper your code but you can not view theirs*. You work needs to be your own. 

If you need to reach out to a CA or the Professor for help, please do so via Piazza and not via email. Answers given via Piazza will help you as well as other students. Thus, emails will always have a lower priority for response than Piazza questions. If you do email the professor or TA, please make a note of what section you are in. This helps us identify you in courseworks faster.

**Multiple Choice**

If the question is multiple choice, you will be given several options to choose from and your function will need to return **one** of those options **verbatim** as a string.

For example: 

Which of the following animals bark?

- dogs
- cats
- fish
- trees

A correct answer would be structured in the following way:

```python
def question_animals():
    return 'dogs'
```

You answer will be stripped of left and right white space and lowercased before comparison to the correct answer during grading.
</details>

<!-- grade_id: cell-79849c5bf45cf5d8, points: 0, grade: true, solution: false -->
```python

```

```python
%pylab inline
```

Output:
```
Populating the interactive namespace from numpy and matplotlib
```

```python

```

**Q**: Write a recursive function `sum_list` to find the sum of all elements in a list. You can assume the list is flat. You can not use the built-in function `sum()`. You can not use iteration via a `for` or `while` loop. Your function must return an integer. If your solution uses exceptions, the except statements must not be bare (ie. the must specify the errors they can catch). You must handle the case where the given list is empty.

[2 points]

<!-- grade_id: cell-1da160cbcad5544d, points: 0, grade: false, solution: true -->
```python
def sum_list(list_):
    print(list_)
    if list_ == []:
        return 0
    return list_[0] + sum_list(list_[1:])
```

```python
##l = [1,2]#np.random.random(300)
#print(l)
sum_list([1,0])
```

Output:
```
[1, 0]
[0]
[]
1
```

<!-- grade_id: cell-ced4e0b96c5fbe9a, points: 0.5, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-923978f1255fae48, points: 0.5, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-601b53d0f1168ba6, points: 1, grade: true, solution: false -->
```python

```

**Q**: Create a recursive function `fib` to find the nth Fibonacci number. Assume that the the sequence starts with `0, 1, 1, ...`. 

```python
>>> fib(0)
0
>>> fib(5)
5
>>> fib(10)
55
```

If your solution uses exceptions, the except statements must not be bare (ie. the must specify the errors they can catch).

[2 points]

<!-- grade_id: cell-be7a8a766b1b4e38, points: 0, grade: false, solution: true -->
```python
solved = {}
def fib(n, depth=0):
    global solved
    if n <=1: 
        return n
    if solved.get(n):
        return solved[n]
    else:
        solved[n] = fib(n-1, depth+1) + fib(n-2, depth+1)
        return fib(n-1, depth+1) + fib(n-2, depth+1)
```

<!-- grade_id: cell-9229e538c5604fb8, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
assert fib(0) == 0
### END TESTS
```

<!-- grade_id: cell-50ad138e1467ce1b, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
assert fib(10) == 55
### END TESTS
```

```python
%%time
fib(3000)
```

Output:
```
CPU times: user 83 µs, sys: 1 µs, total: 84 µs
Wall time: 86.1 µs
410615886307971260333568378719267105220125108637369252408885430926905584274113403731330491660850044560830036835706942274588569362145476502674373045446852160486606292497360503469773453733196887405847255290082049086907512622059054542195889758031109222670849274793859539133318371244795543147611073276240066737934085191731810993201706776838934766764778739502174470268627820918553842225858306408301661862900358266857238210235802504351951472997919676524004784236376453347268364152648346245840573214241419937917242918602639810097866942392015404620153818671425739835074851396421139982713640679581178458198658692285968043243656709796000
```

```python
%%time
x = [None for _ in range(1000000)]
```

Output:
```
CPU times: user 43.7 ms, sys: 4.94 ms, total: 48.6 ms
Wall time: 48 ms
```

```python
((48e-3/1e6)*(2**100))/(525600*60)
```

Output:
```
1929452968383910.8
```

```python

```

<!-- grade_id: cell-e45f33d8783fa401, points: 0.5, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-ef3e9b562e825d7f, points: 0.5, grade: true, solution: false -->
```python

```

**Q**: Write a recursive function `recurse` that determines the recursion limit of Python on your machine. A recursion limit is the maximum number of times a function can call itself. The function should return the recursion limit that is reached. Your function must catch the `RecursionError` error.

NB. The number of times you can call this function is not a constant and it's harder to determine what the number of calls to recurse is without "n". You can ignore this parameter within the body of your function if you wish. However, your function signature must take it.

Assume that the `recurse` function will be called with zero when it is called for the first time.

[2 points]

<!-- grade_id: cell-23cabe3a6fd9d509, points: 0, grade: false, solution: true -->
```python
def recurse(n: int) -> int:
    try:
        return 1 + recurse(n)
    except:
        return 1
```

```python
recurse(0)
```

Output:
```
2964
```

<!-- grade_id: cell-ef5d6fef7fa5b71b, points: 1, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-df21a08199acd722, points: 1, grade: true, solution: false -->
```python

```